Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = ""
COLLABORATORS = ""
```

---

# Preprocessing Files
Different sources and tools may make use of different formats to represent information and output may not directly correspond.
In this course, we will mainly or even exclusively work with the conll format (depending on the preprocessing tools you select later on). Nevertheless, there may be differences in tokenization, in class labels or in the number of columns provided. Depending on what the difference is exactly, you may want to adapt the input files or build in handling mechanisms in the steps that follow (evaluation or feature extraction).
In this case, we are preparing files for evaluation. You can either handle differences as a step in evaluation or make sure data, class labels and structure align during preprocessing. 
We recommend you make sure tokens and class labels align during preprocessing. For structure (fix the column that provides the necessary information, fix the delimiter between columns), you may decide yourself. Note though, that the tests in the previous assignments assume that all files make use of the standard conll delimiter (\t). In this particular case, it is therefore also easier to make sure all files use the same delimiter.

<!-- grade_id: cell-7e753867a9b9b37f, points: 0, grade: false, solution: true -->
```python
import csv

def matching_tokens(conll1, conll2):
    '''
    Check whether the tokens of two conll files are aligned
    
    :param conll1: tokens (or full annotations) from the first conll file
    :param conll2: tokens (or full annotations) from the first conll file
    
    :returns boolean indicating whether tokens match or not
    '''
    # YOUR CODE HERE
    raise NotImplementedError()
    
```

You will find some suggested functions that support the process of comparing and adapting conll files below.
You may want to adapt the arguments and structure (this is not necessary, but feel free if it makes things easier for you). As long as you can run the tests below, this is fine.

<!-- grade_id: cell-6c93b5be48739e09, points: 0, grade: false, solution: true -->
```python
def read_in_conll_file(conll_file, delimiter='\t'):
    '''
    Read in conll file and return structured object
    
    :param conll_file: path to conll_file
    :param delimiter: specifies how columns are separated. Tabs are standard in conll
    
    :returns structured representation of information included in conll file
    '''
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-af3c5a411b7d2103, points: 0, grade: false, solution: true -->
```python
def convert_annotations(conll_object, annotation_identifier, conversions):
    '''
    Check which annotations need to be converted for the output to match and convert them
    
    :param conll_object: structured object with conll annotations
    :param annotation_identifier: indicator of how to find the annotations in the object (e.g. key of dictionary, index of list)
    :param conversions: pointer to the conversions that apply. This can be external (e.g. a local file with conversions) or internal (e.g. prestructured dictionary). In case of an internal object, you probably want to add a function that creates this from a local file.
    
    '''
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-60f9c7a1546da94d, points: 0, grade: false, solution: true -->
```python
def get_predefined_conversions(conversion_file):
    '''
    Read in file with predefined conversions and return structured object that maps old annotation to new annotation
    
    :param conversion_file: path to conversion file
    
    :returns object that maps old annotations to new ones
    '''
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-37479ed77886f209, points: 0, grade: false, solution: true -->
```python
def alignment_okay(conll1, conll2):
    '''
    Read in two conll files and see if their tokens align
    '''
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-5412818440d76f8e, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-9507dcfcd4c75af6, points: 0, grade: false, solution: true -->
```python
def preprocess_files(conll1, conll2, column_identifiers, conversions):
    '''
    Guides the full process of preprocessing files and outputs the modified files.
    
    :param conll1: path to the first conll input file
    :param conll2: path to the second conll input file
    :param column_identifiers: object providing the identifiers for target column
    :param conversions: path to a file that defines conversions
    '''
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-a855818840be7e62, points: 1, grade: true, solution: false -->
```python
from nose.tools import assert_equal

assert_equal(alignment_okay('data/minigold.csv','data/miniout1.csv'),True)
assert_equal(alignment_okay('data/minigold.csv','data/miniout2.csv'),False)

```