Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = ""
COLLABORATORS = ""
```

---

# Verification if programs run from the commandline work
Advanced versions of an assignment and completed components of the final system must run from the commandline.
Autograding for these assignments will call the script with specific settings and check whether expected output is printed to a file. We will test this here through a program that will print "Hello World!" in an output file.
In this case, the hidden text may make use of a different outputfile. In order for verification to work, you can adapt the command_frame, but the call (through os.system) should yield the expected results when other variables change. 

# Specific instructions
The program takes (at least) one argument: the name of the outputfile.
It must print the string "Hello World!" to this outputfile
The setup assumes the program is called 'write_hello_world.py'. Make sure to name your file this way *or* change the command_frame accordingly.

<!-- grade_id: cell-141a12708d713dda, points: 1, grade: true, solution: false -->
```python
import os
from nose.tools import assert_equal

def check_info_in_file(filename, text):
    '''
    checks whether a specific string is present in a file
    :param filename: the name of the verification file
    :param text: the text that is expected
    
    :returns a boolean indicating whether the text is found in the file or not
    '''   
    with open(filename) as f:
        if text in f.read():
            return True
        else:
            return False

    
command_frame = 'python write_hello_world.py'
outputfile = 'helloworld.txt'
os.system(command_frame + ' ' + outputfile)

assert_equal(check_info_in_file(outputfile, 'Hello World!'),True)


```