Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = ""
COLLABORATORS = ""
```

---

# Introduction

The sole purpose of this exercise is to test the NBgrader environment.
You are asked to write two very basic functions and hand in your solution.
Your functions are checked automatically by visible verification (for yourself) and hidden verification (for us).

<!-- grade_id: cell-5c5349ead694717c, points: 0, grade: false, solution: true -->
```python
def provide_sum(x, y):
    '''return the sum of the two input variables'''
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-76148b27cd5fec7d, points: 1, grade: true, solution: false -->
```python
from nose.tools import assert_equal

assert_equal(provide_sum(2,2),4)
assert_equal(provide_sum(22,34),56)

```

<!-- grade_id: cell-5558469c769bbd34, points: 0, grade: false, solution: true -->
```python
def provide_multiplication(x, y):
    '''return the multiplication of the two input variables'''
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-ab4bc6e2375eb1dc, points: 0, grade: true, solution: false -->
```python
assert_equal(provide_multiplication(1,3),3)
assert_equal(provide_multiplication(22,34),748)
```