Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = ""
COLLABORATORS = ""
```

---

# Basic Evaluation
Create a set-up that can determine precision recall and f-score of the system output of multiple systems compared to a gold standard.
This component should work with the output of your preprocessing component.
It should be possible to call the evaluation script by specifing a gold file and system output files as input, together with any information necessary to correctly process the files.
The file will also provide tables in latex form as output. You can choose whether you want to write this to a default outputfile or whether you want to ask for an (obligatory) argument specifying the outputfile.

For verification reasons, the program should also output a file called 'evaluation_outcome.txt' that provides the outcome of the evaluation on individual lines. e.g.:
system1 I-ORG precision 0.623

<!-- grade_id: cell-b6cce6ab0e0b2c3e, points: 0, grade: true, solution: false -->
```python
import os
from nose.tools import assert_equal


def check_info_in_file(filename, text):
    '''
    checks whether a specific string is present in a file
    :param filename: the name of the verification file
    :param text: the text that is expected
    
    :returns a boolean indicating whether the text is found in the file or not
    '''   
    with open(filename) as f:
        if text in f.read():
            return True
        else:
            return False

#set command to run your program the second test should require the same system info as the first (if applicable)

###YOU CAN CHANGE THIS COMMAND BELOW IF YOU HAVE TO
command = 'python evaluation.py data/minigold.csv'

###THIS IS A PLACE HOLDER FOR THE NAME OF YOUR SYSTEM. (HERE MINIOUT1, OTHERWISE E.G. SPACY/STANFORD)
systemname = 'miniout1'

###THIS IS WHERE YOU SPECIFY NEEDED INFORMATION ON THE FILE YOU WANT TO TEST. HERE IT IS JUST THE FILE NAME ITSELF
systeminfo = 'data/miniout1.csv'
filename = 'evaluation_outcome.txt'
text = 'miniout1 O f-score 0.889'

#hint: if the variable check_eval stores the f-score, you obtain 3 decimals using "%.3f" % check_eval


```

<!-- grade_id: cell-03fccf575b67e79e, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-744abb4dde824eb9, points: 1, grade: true, solution: false -->
```python


os.system(command + systemname + systeminfo)
assert_equal(check_info_in_file(filename, text),True)
```