Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = ""
COLLABORATORS = ""
```

---

# Basic Evaluation
Create a set-up that can determine precision recall and f-score of the system output of multiple systems compared to a gold standard.
This component should work with the output of your preprocessing component.
It should be possible to call the evaluation script by specifing a gold file and system output files as input, together with any information necessary to correctly process the files.
The file will also provide tables in latex form as output. You can choose whether you want to write this to a default outputfile or whether you want to ask for an (obligatory) argument specifying the outputfile.

For verification reasons, the program should also output a file called 'evaluation_outcome.txt' that provides the outcome of the evaluation on individual lines. e.g.:
system1 I-ORG precision 0.623

<!-- grade_id: cell-b6cce6ab0e0b2c3e, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-744abb4dde824eb9, points: 1, grade: true, solution: false -->
```python

assert_equal(check_info_in_file(filename, text),True)
```