Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = ""
COLLABORATORS = ""
```

---

# Instructions
Write a program that takes care of preprocessing of conll files. The code should:
1. verify whether the tokens of two conll files align
2. convert annotations so that both files make use of the same conventions (if applicable)

The program should output a log file called preprocessing.log that specifies which files corresponded and which files did not. This should be formatted in the following way (with filename1 and filename2), depending on which applies:
filename1 and filename2: ALIGN
filename1 and filename2: MISSMATCH

The log should furthermore specify whether a specific file contains converted annotations:
filename2 contains CONVERSIONS

<!-- grade_id: cell-304f99557b9b6d53, points: 0, grade: true, solution: false -->
```python
import os
from nose.tools import assert_equal


def check_info_in_file(filename, text):
    '''
    checks whether a specific string is present in a file
    :param filename: the name of the verification file
    :param text: the text that is expected
    
    :returns a boolean indicating whether the text is found in the file or not
    '''   
    with open(filename) as f:
        if text in f.read():
            return True
        else:
            return False

command_frame = 'python preprocess.py '
logfile = 'preprocessing.log'

goldfile = 'data/minigold.csv'
systemout = 'data/miniout1.csv'
os.system(command_frame + goldfile + ' ' + systemout)

assert_equal(check_info_in_file(logfile, goldfile + ' ' + systemout + ': ALIGN'),True)

```