Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = ""
COLLABORATORS = ""
```

---

# Instructions
Write a program that takes care of preprocessing of conll files. The code should:
1. verify whether the tokens of two conll files align
2. convert annotations so that both files make use of the same conventions (if applicable)

The program should output a log file called preprocessing.log that specifies which files corresponded and which files did not. This should be formatted in the following way (with filename1 and filename2), depending on which applies:
filename1 and filename2: ALIGN
filename1 and filename2: MISSMATCH

The log should furthermore specify whether a specific file contains converted annotations:
filename2 contains CONVERSIONS

<!-- grade_id: cell-304f99557b9b6d53, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-97898ef964a3640a, points: 1, grade: true, solution: false -->
```python
assert_equal(check_info_in_file(logfile, goldfile + ' ' + systemout + ': ALIGN'),True)
```