# Instructions
Write a program that takes care of preprocessing of conll files. The code should:
1. verify whether the tokens of two conll files align
2. convert annotations so that both files make use of the same conventions (if applicable)

The program should output a log file called preprocessing.log that specifies which files corresponded and which files did not. This should be formatted in the following way (with filename1 and filename2), depending on which applies:
filename1 and filename2: ALIGN
filename1 and filename2: MISSMATCH

The log should furthermore specify whether a specific file contains converted annotations:
filename2 contains CONVERSIONS

<!-- grade_id: cell-304f99557b9b6d53, points: 0, grade: false, solution: true -->
```python
import os
from nose.tools import assert_equal


def check_info_in_file(filename, text):
    '''
    checks whether a specific string is present in a file
    :param filename: the name of the verification file
    :param text: the text that is expected
    
    :returns a boolean indicating whether the text is found in the file or not
    '''   
    with open(filename) as f:
        if text in f.read():
            return True
        else:
            return False

##SAMPLE COMMAND FRAME: CHANGE IF YOU GAVE YOUR PROGRAM A DIFFERENT NAME
command_frame = 'python preprocess.py '

#PLEASE KEEP THIS STABLE
logfile = 'preprocessing.log'

#THESE WILL BE DEFINED AS PART OF THE TEST
goldfile = 'data/minigold.csv'

#THESE WILL BE DEFINED AS PART OF THE TEST
systemout = 'data/miniout1.csv'



```

<!-- grade_id: cell-fa3a20c5fe8f37ee, points: 0, grade: false, solution: true -->
```python
###REDEFINE COMMAND HERE IF YOU NEED TO
# command_frame = ''
```

<!-- grade_id: cell-97898ef964a3640a, points: 1, grade: true, solution: false -->
```python

os.system(command_frame + goldfile + ' ' + systemout)
assert_equal(check_info_in_file(logfile, goldfile + ' ' + systemout + ': ALIGN'),True)
### BEGIN HIDDEN TESTS 

goldfile = '../../testdata/minigold.csv'
systemout = '../../testdata/miniout2.csv'
os.system(command_frame + goldfile + ' ' + systemout)

assert_equal(check_info_in_file(logfile, goldfile + ' ' + systemout + ': MISMATCH'),True)
assert_equal(check_info_in_file(logfile, goldfile + ' ' + systemout + ': ALIGN'),False)
### END HIDDEN TESTS
```